// Numbas version: exam_results_page_options
{"name": "Forma reducida por filas", "extensions": ["/srv/numbas/media/user-extensions/extracted/224/linear-algebra"], "custom_part_types": [], "resources": [], "navigation": {"allowregen": true, "showfrontpage": false, "preventleave": false}, "question_groups": [{"pickingStrategy": "all-ordered", "questions": [{"name": "Forma reducida por filas", "tags": [], "metadata": {"description": "<p>C&aacute;lculo de forma normal reducida por filas.</p>", "licence": "Creative Commons Attribution 4.0 International"}, "statement": "<p>Calcula la forma normal reducida por filas $H$ de la matriz \\(A=\\var{A}\\).</p>\n<p></p>", "advice": "<p>Sigue las indicaciones que se muestran <a href=\"https://www.ugr.es/~lmerino/1-1.html\" target=\"_blank\" rel=\"noopener\" title=\"forma normal de Hermite por filas\">aqu&iacute;</a>.</p>", "rulesets": {}, "extensions": ["linear-algebra"], "builtin_constants": {"e": true, "pi,\u03c0": true, "i": true}, "constants": [], "variables": {"A": {"name": "A", "group": "Ungrouped variables", "definition": "matrix(repeat(repeat(random(-3..3),4),3))", "description": "", "templateType": "anything", "can_override": false}, "rA": {"name": "rA", "group": "Ungrouped variables", "definition": "reduced_row_echelon_form(A)", "description": "<p>forma reducida por filas de A</p>", "templateType": "anything", "can_override": false}}, "variablesTest": {"condition": "", "maxRuns": 100}, "ungrouped_variables": ["A", "rA"], "variable_groups": [], "functions": {}, "preamble": {"js": "", "css": ""}, "parts": [{"type": "matrix", "useCustomName": true, "customName": "Matriz reducida", "marks": 1, "scripts": {}, "customMarkingAlgorithm": "", "extendBaseMarkingAlgorithm": true, "unitTests": [], "showCorrectAnswer": true, "showFeedbackIcon": true, "variableReplacements": [], "variableReplacementStrategy": "originalfirst", "nextParts": [], "suggestGoingBack": false, "adaptiveMarkingPenalty": 0, "exploreObjective": null, "prompt": "<p>Introduce las entradas de $H$.</p>", "correctAnswer": "rA", "correctAnswerFractions": true, "numRows": "3", "numColumns": "4", "allowResize": false, "tolerance": 0, "markPerCell": false, "allowFractions": true, "minColumns": 1, "maxColumns": 0, "minRows": 1, "maxRows": 0, "prefilledCells": ""}], "partsMode": "all", "maxMarks": 0, "objectives": [], "penalties": [], "objectiveVisibility": "always", "penaltyVisibility": "always", "contributors": [{"name": "Pedro A. Garc\u00eda-S\u00e1nchez", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1217/"}]}]}], "contributors": [{"name": "Pedro A. Garc\u00eda-S\u00e1nchez", "profile_url": "https://numbas.mathcentre.ac.uk/accounts/profile/1217/"}]}